# Model Selection and Evaluation - Part 1

When building supervised machine learning models, we need to solve two problems:

1. **Model selection** - Finding the model that does as well as possible
on our learning task.

2. **Model evaluation** - Predicting **generalization error**, or the expected performance of our
model on unseen data.

Both are critical.  Without 1. we can't have an effective model and without 2. we can't *know* if we have an effective model.

Once we have picked a particular machine learning algorithm, model
selection comes down to the problem of **hyperparameter** tuning.
Hyperparameters are parameters of our learning models that need to be
selected before the model can be learned.  Examples include the
maximum number of leaves in a decision tree or the number of hidden
units in a neural network classifier.  The **parameters** of our model
are learning from the training data, for example, the feature and split point
to use to for a node in a decision tree.  Parameters are learned from the training data,
and the hyperparameters guide the learning process.  

**WARNING:**  The below activities showcase several *BAD* approaches to model selection and evaluation.  These examples are not meant to illustrate the correct way of doing things, they are meant to show the consequences of doing things incorrectly. 


## Activity 1 - Naive Model Selection

For now, let's focus entirely on model selection and disregard model
evaluation.  The following cell will load a data set and use a
decision tree classifier to fit a decision tree to the data. Try adjusting the `max_leaf_nodes` hyperparameter in order to minimize the error rate on the training set.

```python
import numpy as np
import matplotlib.pyplot as plt
from DTClassBounds import decision_areas,plot_areas

from sklearn.datasets import make_moons
from sklearn.tree import DecisionTreeClassifier
from matplotlib.colors import ListedColormap

X, y = make_moons(n_samples=200, noise=0.3, random_state=0)
# Build a decision tree regressor

tree = DecisionTreeClassifier(max_leaf_nodes=27)
tree.fit(X, y)

# Evaluate the error rate of our decision tree on the data set
y_predict = tree.predict(X)
error_rate = np.sum((y - y_predict)**2) / y.size
print("error rate: {:.4f}".format(error_rate))

# Visualize the decision boundaries
cm = plt.cm.RdBu
cm_bright = ListedColormap(['#FF0000', '#0000FF'])
plt.scatter(X[:,0], X[:,1], c=y, cmap=cm_bright)
rec = decision_areas(tree, [-2,3,-2,3])
plot_areas(rec)
plt.show()

```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
c:\Users\tommy\Desktop\CS445\lab2\model_selection_classification_1.ipynb Cell 3 in <cell line: 1>()
----> <a href='vscode-notebook-cell:/c%3A/Users/tommy/Desktop/CS445/lab2/model_selection_classification_1.ipynb#W2sZmlsZQ%3D%3D?line=0'>1</a> import numpy as np
      <a href='vscode-notebook-cell:/c%3A/Users/tommy/Desktop/CS445/lab2/model_selection_classification_1.ipynb#W2sZmlsZQ%3D%3D?line=1'>2</a> import matplotlib.pyplot as plt
      <a href='vscode-notebook-cell:/c%3A/Users/tommy/Desktop/CS445/lab2/model_selection_classification_1.ipynb#W2sZmlsZQ%3D%3D?line=2'>3</a> from DTClassBounds import decision_areas,plot_areas

ModuleNotFoundError: No module named 'numpy'
```

# <span style="color:red">Question #1</span>

* What value of the hyperparameter resulted in the lowest error rate?

<!-- grade_id: cell-bfa3abddf6a65c1c, points: 5, grade: true, solution: true -->
## *YOUR ANSWER HERE*

# <span style="color:red">Question #2</span>

Do you think that this error rate reflects how well this model will do on unseen data?  Why or why not?

<!-- grade_id: cell-7edd77275ac77e02, points: 0, grade: true, solution: true -->
## *YOUR ANSWER HERE*

## New Data Arrives!

In the exercise above, you tuned the hyperparameters to achieve the *best* fit the training data.  Now let's see what happens when we use this model on some new data drawn from the same underlying distribution:

```python
X_test, y_test = make_moons(n_samples=200, noise=0.3, random_state=1)

print('Using your best model, which was fit in a prior cell with {} leaves'
     .format(tree.get_n_leaves()))

y_test_predict = tree.predict(X_test)
error_rate = np.sum((y_test - y_test_predict)**2) / y_test.size

print("error rate: {:.4f}".format(error_rate))
```

## Activity 2 - Using a Test Set for Hyperparameter Tuning and Evaluation

In the exercise above, you were able to perfectly fit a training data set, but that didn't tell you anything about how well your model would perform on unseen data.  Recall that we want our models to **generalize**, that is, perform well on data that the model has not seen
previously. 

One approach to address this is by splitting our limited data into a **training set** and a **test** set.  This is illustrated in the cell below.

```python
X, y = make_moons(n_samples=200, noise=0.3, random_state=0)

# Split our data into a training and testing set...
split_point = int(X.shape[0] * .8) # Use 80% of the data to train the model

X_train = X[0:split_point, :]
y_train = y[0:split_point]

X_test = X[split_point::, :]
y_test = y[split_point::]

# Build a decision tree regressor using the TRAINING set
tree = DecisionTreeClassifier(max_leaf_nodes=3)
tree.fit(X_train, y_train)

# Evaluate the error rate of our decision tree on the TESTING set 
y_test_predict = tree.predict(X_test)
error_rate= np.sum((y_test - y_test_predict)**2) / y_test.size

print("test data error rate: {:.4f}".format(error_rate))
```

# <span style="color:red">Student Activity (Question #3)</span>

In the below code block, write a loop that iterates over the hyperparameter for
the number of leaves and reports the value that produces the lowest error rate
on the **test** data.  

Since this class focuses on visualization and explaining the results,
for each hyperparameter setting, store the error rate for
the training and test sets in 2 separate python lists.  You should be
able to deduce the number of leaves used by the position in the list
(the first position used 2 leaves, the second position in the list used
3 leaves, ...).



<!-- grade_id: cell-fd99b32e606cdaeb, points: 0, grade: false, solution: true -->
```python
## Answer to question 3 goes here

X, y = make_moons(n_samples=200, noise=0.3, random_state=0)

# Split our data into a training and testing set...
split_point = int(X.shape[0] * .8) # Use 80% of the data to train the model

X_train = X[0:split_point, :]
y_train = y[0:split_point]

X_test = X[split_point::, :]
y_test = y[split_point::]

print(X_train.shape)

"""
Code a loop that varies the hyperparameter for the
number of leaves in the tree and records BOTH
the error rate on the training data and the 
error rate on the test data.  Vary the number
of leaves from 2 to the number of data points.
"""

train_error_rates = [0] * (X_train.shape[0] - 1)
test_error_rates =  [0] * (X_train.shape[0] - 1)

# YOUR CODE HERE
raise NotImplementedError()

## Find the best value (NOT using a for loop)
## The BEST is the lowest test error rate with the fewest leaves

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8a0510f094466364, points: 5, grade: true, solution: false -->
```python
# This is code to test your answer to Question 3
# just run it after running your code and see how your answers compares
assert(len(train_error_rates) == (X_train.shape[0] -1))

# test the first and last values 
assert(train_error_rates[0] == 0.1875)
assert(test_error_rates[0] == 0.15)

assert(train_error_rates[-1] == 0.0)
assert(test_error_rates[-1] == 0.175)


train_error_rates_ref = [0.1875, 0.1375, 0.1125, 0.1125, 0.1125, 0.08125, 0.08125,
                     0.06875, 0.0625, 0.05625, 0.05, 0.05, 0.05, 0.05, 0.0375,
                     0.03125, 0.025, 0.025, 0.0125, 0.0125, 0.00625, 0.00625, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 
                     0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0]

print(int(len(train_error_rates_ref)))

test_error_rates_ref = [0.15, 0.1, 0.075, 0.075, 0.075, 0.175, 0.175, 0.175, 0.075, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 0.175, 
                        0.175, 0.175, 0.175, 0.175, 0.175]

np.testing.assert_array_equal(train_error_rates_ref, train_error_rates,
                                      err_msg='Test error rate array mismatch')
np.testing.assert_array_equal(test_error_rates_ref, test_error_rates,
                                      err_msg='Test error rate array mismatch')
```

# <span style="color:red">Question #4</span>

- What hyperparameter settings gives us the lowest error rate on the **test** data?  Provide both the hyperparameter value and the test error rate?   


<!-- grade_id: cell-6556a7e8842e725c, points: 5, grade: true, solution: true -->
## *YOUR ANSWER HERE*

# <span style="color:red">Question 5</span>

* Create a plot with number of leaves vs. error rates (both training and test) varying the maximum number of leaves from 2 to 25.

* Do you think *this* error rate will be reflective of how well our model will perform on unseen data? 

<!-- grade_id: cell-954338e7a6b2069d, points: 0, grade: true, solution: true -->
```python
"""
Write your code to generate your plot here for Question 5
"""

# YOUR CODE HERE
raise NotImplementedError()
```

# <span style="color:red">Question 6</span>

* Do you think *this* error rate will be reflective of how well our model will perform on unseen data? 

<!-- grade_id: cell-c1f723fd33396e77, points: 5, grade: true, solution: true -->
## *YOUR ANSWER HERE*

# A Danger in ML -- Model Overfitting

Discussed in IDD Chapter 3.4.

Looking at the model's error_rate over both the training set and testing set as a function of the maximum number of leaves offers a lot of insight.

When the number of leaves is low (2 for example), we get a very simple model (can't get simplier than two leaves).  When both the training and test error rates are high, we characerize the model as **underfitting** the data.  The model simply does not have enough parameters to capture the relationship between the features and the class labels. 

When the number of leaves is very high, we get a complex model.  The tree for this model has 
many decisions/branches, and in the extreme case, can have leaves with one a single example
from the training set.  So, you might ask, how can a model generalize if the leaves 
don't contain more than one example?  The answer in general, is that **is can not**.

One way of the thinking about a model like this is that it is memorizing the
answers instead of learning the relationship between the features and the class labels.  This becomes
apparent when, given new data (the test set), the model does not perform as well as on the training
data.  This phenomena is known as **model overfitting** and is a very important concept, as it has
lead to some of the worst failures in the application of machine learning.  

The tool we have at this point to identify underfitting and overfitting is the plot 
of the error rate on the training and testing datasets.

```python

```

# Expected Generalization Error 

Notice that in this example we are using our test set for *both* model selection and model evaluation.  We used it for **model selection** by searching for a hyperparameter setting that minimizes error on the test set.  We use it for **model evaluation** by using our test set error as an estimate of the expected error rate on unobserved data.

Let's see how our model does on some new, unobserved data drawn from the same distribution.  This cell will give us a good estimate of our *actual* generalization error.  (Note that in real-world problems we can't run a test like this because we don't have unlimited access to extra data that we can use to check our work.)

```python
tree = DecisionTreeClassifier(max_leaf_nodes=????) # Put your best hyperparameter here!
tree.fit(X_train, y_train)

# Let's see how we do on unobserved data... 
X_new, y_new = make_moons(n_samples=200, noise=0.3, random_state=2)
y_new_predict = tree.predict(X_new)
error_rate = np.sum((y_new - y_new_predict)**2) / y_new.size
print("Error rate: {:.4f}".format(error_rate))
```

# <span style="color:red">Question 7</span>

* Relative to Activity 1, where we just looked for the model that best fit our training data, would you say that our train/test split was beneficial in terms of model selection, i.e. did we end up with a better model? Justify your answer.


# <span style="color:red">Answer to 7</span>
Your answer goes here


# <span style="color:red">Question 8</span>
* Would you say that our train/test split was beneficial in terms of model evaluation, i.e. were we able make a better prediction of our generalization error?  Justify your answer.

<!-- grade_id: cell-d024e59d62785d9f, points: 5, grade: true, solution: true -->
## *YOUR ANSWER HERE*

# <span style="color:red">Question 9</span>  

Do you see any problems here in terms of model selection or evaluation?  How accurate was our prediction of generalization error?

<!-- grade_id: cell-faaa824dd742dd94, points: 5, grade: true, solution: true -->
## *YOUR ANSWER HERE*

### Click [here](model_selection_classification_2.ipynb) to open the next page of exercises...

```python

```